# Part 2: Mining the web: Web APIs

We hope the preceding exercise was painful: even with tools to process HTML, it is rough downloading raw HTML and trying to extract information from it!

> Can you think of any other reasons why scraping websites for data in this way is not a good idea?

Luckily, many websites provide an application programming interface (API) for querying their data or otherwise accessing their services from your programs. For instance, Twitter provides a web API for gathering tweets, Flickr provides one for gathering image data, and Github for accessing information about repository histories.

These kinds of web APIs are much easier to use than, for instance, the preceding technique which scrapes raw web pages and then has to parse the resulting HTML. Moreover, there are more scalable in the sense that the web servers can transmit structured data in a less verbose form than raw HTML.

As a starting example, here is some code to look at the activity on Github related to the public version of our course's materials.

```python
import requests

response = requests.get ('https://api.github.com/repos/cse6040/labs-fa17/events')

print ("==> .headers:", response.headers, "\n")
```

Output:
```
==> .headers: {'Server': 'GitHub.com', 'Date': 'Fri, 15 Feb 2019 21:19:41 GMT', 'Content-Type': 'application/json; charset=utf-8', 'Transfer-Encoding': 'chunked', 'Status': '200 OK', 'X-RateLimit-Limit': '60', 'X-RateLimit-Remaining': '26', 'X-RateLimit-Reset': '1550266846', 'Cache-Control': 'public, max-age=60, s-maxage=60', 'Vary': 'Accept', 'ETag': 'W/"75e5e4b88e11487c827e95f98d8d1b3b"', 'Last-Modified': 'Tue, 05 Feb 2019 14:27:47 GMT', 'X-Poll-Interval': '60', 'X-GitHub-Media-Type': 'github.v3; format=json', 'Access-Control-Expose-Headers': 'ETag, Link, Location, Retry-After, X-GitHub-OTP, X-RateLimit-Limit, X-RateLimit-Remaining, X-RateLimit-Reset, X-OAuth-Scopes, X-Accepted-OAuth-Scopes, X-Poll-Interval, X-GitHub-Media-Type', 'Access-Control-Allow-Origin': '*', 'Strict-Transport-Security': 'max-age=31536000; includeSubdomains; preload', 'X-Frame-Options': 'deny', 'X-Content-Type-Options': 'nosniff', 'X-XSS-Protection': '1; mode=block', 'Referrer-Policy': 'origin-when-cross-origin
```

Note the `Content-Type` of the response:

```python
print (response.headers['Content-Type'])
```

Output:
```
application/json; charset=utf-8
```

The response is in JSON format, which is an open format for exchanging semi-structured data. (JSON stands for **J**ava**S**cript **O**bject **N**otation.) JSON is designed to be human-readable and machine-readable, and maps especially well in Python to nested dictionaries. Let's take a look.

> See also [this tutorial](https://www.w3schools.com/js/js_json_intro.asp) for a JSON primer. JSON is among _the_ universal formats for sharing data on the web; see, for instance, https://www.sitepoint.com/10-example-json-files/.

```python
import json
print(type(response.json ()))
print(json.dumps(response.json()[:3], sort_keys=True, indent=2))
```

Output:
```
<class 'list'>
[
  {
    "actor": {
      "avatar_url": "https://avatars.githubusercontent.com/u/31219281?",
      "display_login": "vorsatti",
      "gravatar_id": "",
      "id": 31219281,
      "login": "vorsatti",
      "url": "https://api.github.com/users/vorsatti"
    },
    "created_at": "2019-02-05T14:27:47Z",
    "id": "9011334282",
    "org": {
      "avatar_url": "https://avatars.githubusercontent.com/u/31073927?",
      "gravatar_id": "",
      "id": 31073927,
      "login": "cse6040",
      "url": "https://api.github.com/orgs/cse6040"
    },
    "payload": {
      "action": "started"
    },
    "public": true,
    "repo": {
      "id": 100506580,
      "name": "cse6040/labs-fa17",
      "url": "https://api.github.com/repos/cse6040/labs-fa17"
    },
    "type": "WatchEvent"
  },
  {
    "actor": {
      "avatar_url": "https://avatars.githubusercontent.com/u/10697135?",
      "display_login": "rajpaz",
      "gravatar_id": "",
      "id": 10697135,
      "login": "rajpaz",
 
```

**Exercise 0.** It should be self-evident that the JSON response above consists of a sequence of records, which we will refer to as _events_. Each event is associated with an _actor_. Write some code to extract a dictionary of all actors, where the key is the actor's login and the value is the actor's URL.

<!-- grade_id: extract_actors, points: 0, grade: false, solution: true -->
```python
def extract_actors (json_github_events):
    """Given JSON records for events in a GitHub repo,
    returns a dictionary of the actors and their URLs.
    """
    actors = {}
    for e in json_github_events:
        actors[e['actor']['display_login']] = e['actor']['url']
    return actors
```

<!-- grade_id: extract_actors_test, points: 0, grade: true, solution: false -->
```python
actor_urls = extract_actors(response.json ())

for actor, url in actor_urls.items ():
    print ('{}: {}'.format(actor, url))
    assert url == "https://api.github.com/users/{}".format(actor)
```

Output:
```
vorsatti: https://api.github.com/users/vorsatti
rajpaz: https://api.github.com/users/rajpaz
jxzhou131: https://api.github.com/users/jxzhou131
MCBoarder289: https://api.github.com/users/MCBoarder289
djbermejo: https://api.github.com/users/djbermejo
keithdowd: https://api.github.com/users/keithdowd
schr0841: https://api.github.com/users/schr0841
sanation: https://api.github.com/users/sanation
ddthomas26: https://api.github.com/users/ddthomas26
```

**Exercise 1.** Write some code that goes to each actor's URL and determines their name. If an actor URL is invalid, that actor should not appear in the output.

<!-- grade_id: get_names, points: 0, grade: false, solution: true -->
```python
def lookup_names (actor_urls):
    """Given a dictionary of (actor, url) pairs, looks up the JSON at
    the URL and extracts the user's name (if any). Returns a new
    dictionary of (actor, name) pairs.
    """
    names = {}
    for actor, url in actor_urls.items():
        response = requests.get(url)
        names[actor] = response.json()['name']
    return names
```

<!-- grade_id: get_names_test, points: 0, grade: true, solution: false -->
```python
actor_names = lookup_names (actor_urls)

for actor, name in actor_names.items ():
    print ("{}: {}".format (actor, name))
    
assert actor_names['vorsatti'] == 'Valerio Orsatti'
```

Output:
```
vorsatti: Valerio Orsatti
rajpaz: None
jxzhou131: jxzhou
MCBoarder289: Michael Chapman
djbermejo: None
keithdowd: Keith
schr0841: Greg Schreiter
sanation: None
ddthomas26: None
```

That's the end of this notebook. Processing JSON is fairly straightforward, because it maps very naturally to nested dictionaries in Python. You might search the web for other sources of JSON data, including [this one](https://www.yelp.com/dataset/challenge), and do your own processing!